# EEN060 - Applied object-oriented programming

Teacher: [Carlos Natalino](https://www.chalmers.se/en/persons/carda/) / Examiner: [Paolo Monti](https://www.chalmers.se/en/persons/mpaolo/)

[Canvas course page](https://chalmers.instructure.com/courses/28094)

[Course channel on Chalmers Play](https://play.chalmers.se/channel/EEN060_EEN065_Applied_object-oriented_programming/300149)

Before you turn this assignment list in, make sure everything runs as expected.
First, **restart the kernel** and then **run all cells**.
Then, check if all the tests run correctly.
Note that if one of the problems present an error, the following ones **will not** be tested.

In case of discrepancies between the problem command and the tests, you should solve it having in mind the tests.

There are two types of cell:
1. *solution cells:* These are the cells where you write your answer, or modify the existing code to solve the problem.
2. *test cells:* These cells are used to test whether your solution is correct or not. If the tests run correctly, you should see a message `tests passed`. Otherwise, you should see an error message.

**Delete** the line `raise NotImplementedError()` from the problems that you solve.

**Do not delete or add any cell in this file.** All cells that you need are already in place.

You may also have problems that will be manually graded (e.g., HTML/CSS/JS/pseudocode problems).
These problems **do not** have a respective *test cell*.

If you want to execute a cell, select the cell and press **CTRL+Enter** (in Windows) or **CMD+Enter** (in macOS) or click on the **Run cell** button.

**If you want, you can solve this programming assignment using Google Colab**

Link: https://colab.research.google.com/

Just copy the following line to a code line and run it.

# Assignment 1

In this assignment, students should practice:
- how to use Visual Studio Code
- how to download, solve and submit programming assignments
- how to create functions
- start getting familiar with Python
- how to manipulate variables
- how to use math operations
- how to output the results
- understand different types of variables
- understand the code quality validation

**Observations:**
For the problems in this assignment, you should:
- use only the math operations studied in the lecture notes (i.e., no advanced functions allowed).
- not use external modules, i.e., `import` command is not allowed in your solutions (with exception of the module `typing`).
- make sure that your `utils.py` file has the following line (i.e., replace `onu2` by `onu1` if needed):

`url_python_validator = "https://onu1.s2.chalmers.se/pnu/"`

**Preparation:** Run the cell below every time you start working on this file, and every time you restart the kernel.

```python
%load_ext autoreload
try:
    from utils import validate_python_code, validate_html
except:
    import sys
    print("It seems this file is in the wrong folder. "
          "Make sure to place it in the `programming-assignments` folder/project.",
          file=sys.stderr)
```

Output:
```
The autoreload extension is already loaded. To reload it, use:
  %reload_ext autoreload
```

## Level 1

Level 1 problems are simple problems that require only the knowledge obtained from the lecture's content and the lecture notes.

### Problem 1.1

Create a Python function called `sum_of_digits` that receives a positive integer upto three-digits and returns the sum of the digits.

```
Write your pseudocode

This is not graded

This will not affect your Python code
```

<!-- grade_id: cell-answer-1-1, points: 0, grade: false, solution: true -->
```python
%%writefile sum_of_digits_solution.py
# solution cell
### BEGIN SOLUTION


def sum_of_digits(x: int) -> int:
    a: int = x // 100
    b: int = (x - a * 100) // 10
    c: int = x - a * 100 - b * 10
    return a + b + c


### END SOLUTION
```

Output:
```
Overwriting sum_of_digits_solution.py
```

```python
%autoreload 2
# uncomment the line(s) below to debug
from sum_of_digits_solution import sum_of_digits


a = 93
r = sum_of_digits(a)
print(f"The sum of digits of {a} is {r}")
print("Execution finished", "\u2713")
```

Output:
```
The sum of digits of 93 is 12
Execution finished ✓
```

<!-- grade_id: cell-tests-1-1, points: 1, grade: true, solution: false -->
```python
%autoreload 2
# test cell
try:
    import sum_of_digits_solution
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from sum_of_digits_solution import sum_of_digits
except:
    raise ValueError("The module does not have the necessary function!")

from nose.tools import assert_equal, assert_raises

assert_equal(sum_of_digits(103), 4)
assert_equal(sum_of_digits(123), 6)
assert_equal(sum_of_digits(67), 13)
assert_raises(TypeError, sum_of_digits, "2")
assert_raises(TypeError, sum_of_digits, None)
assert_raises(TypeError, sum_of_digits, "Test", 2)

from utils import validate_python_code

validate_python_code("sum_of_digits_solution.py")
print("tests passed", "\u2713")
```

Output:
```
<IPython.core.display.HTML object>
tests passed ✓
```

### Problem 1.2

Create a Python function called `print_sum_of_digits` that receives a positive integer upto three-digits and print a $\underline{\textbf{string}}$ representing the sum of the digits. Note that the output is returned as a string rather than an integer or float.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-1-2, points: 0, grade: false, solution: true -->
```python
%%writefile print_sum_of_digits_solution.py
# solution cell
### BEGIN SOLUTION


def print_sum_of_digits(x: int) -> str:
    a: int = x // 100
    b: int = (x - a * 100) // 10
    c: int = x - a * 100 - b * 10
    y: str = str(a + b + c)
    return print(y)


### END SOLUTION

```

Output:
```
Overwriting print_sum_of_digits_solution.py
```

```python
%autoreload 2
# uncomment the line(s) below to debug
from print_sum_of_digits_solution import print_sum_of_digits


a = 93
print(f"The sum of digits of {a} is ")
print_sum_of_digits(a)
print("Execution finished", "\u2713")
```

Output:
```
The sum of digits of 93 is 
12
Execution finished ✓
```

<!-- grade_id: cell-test-1-2, points: 1, grade: true, solution: false -->
```python
%autoreload 2
# test cell
try:
    import print_sum_of_digits_solution
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from print_sum_of_digits_solution import print_sum_of_digits
except:
    raise ValueError("The module does not have the necessary function!")

from nose.tools import assert_equal, assert_raises
from unittest.mock import patch
from io import StringIO

# first test
with patch('sys.stdout', new=StringIO()) as fakeOutput:
    print_sum_of_digits(81)
    assert_equal(fakeOutput.getvalue().strip(), '9')

# second test
with patch('sys.stdout', new=StringIO()) as fakeOutput:
    print_sum_of_digits(103)
    assert_equal(fakeOutput.getvalue().strip(), '4')

# testing types
assert_raises(TypeError, print_sum_of_digits)
assert_raises(TypeError, print_sum_of_digits, "2")
assert_raises(TypeError, print_sum_of_digits, 2, "Test")

from utils import validate_python_code

validate_python_code("print_sum_of_digits_solution.py")
print("tests passed", "\u2713")
```

Output:
```
<IPython.core.display.HTML object>
Cell In[168], line 33
     29 assert_raises(TypeError, print_sum_of_digits, 2, "Test")
     31 from utils import validate_python_code
---> 33 validate_python_code("print_sum_of_digits_solution.py")
     34 print("tests passed", "\u2713")

File d:\System\Python_assignments\programming-assignments\utils.py:139, in validate_python_code(filename, **args)
    134     raise ValueError(
    135         f"Error while decoding the response. "
    136         f"Detail: {e}. Response: {response_html.text}"
    137     )
    138 if data["fail"]:
--> 139     raise ValueError("The code needs change. Check messages above.")

ValueError: The code needs change. Check messages above.
```

### Problem 1.3

Create a function called `sphere_volume` that, based on one parameter representing the radius of the sphere (integer or float), prints its volume upto three decimal places.

Remember that the volume $V$ of a sphere given its radius $r$ is given by:
\begin{equation*}
V = \dfrac{4}{3}\pi r^3. 
\end{equation*}

**Observations:**
- You should assume $\pi=3.14$.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-1-3, points: 0, grade: false, solution: true -->
```python
%%writefile sphere_volume_solution.py
# solution cell
### BEGIN SOLUTION


def sphere_volume(x: int | float) -> str:
    v: int | float = (4 / 3) * 3.14 * (x ** 3)
    return round(v, 3)


### END SOLUTION

```

Output:
```
Overwriting sphere_volume_solution.py
```

```python
%autoreload 2
# uncomment the line(s) below to debug
from sphere_volume_solution import sphere_volume


r = 0.11
print(f'The volume of a sphere with radius {r} is {sphere_volume(r)}')
sphere_volume(r)
print("Execution finished", "\u2713")
```

Output:
```
The volume of a sphere with radius 0.11 is 0.006
Execution finished ✓
```

<!-- grade_id: cell-test-1-3, points: 1, grade: true, solution: false -->
```python
%autoreload 2
# test cell
try:
    import sphere_volume_solution
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from sphere_volume_solution import sphere_volume
except:
    raise ValueError("The module does not have the necessary function!")
    
from nose.tools import assert_equal, assert_raises
from unittest.mock import patch
from io import StringIO

with patch('sys.stdout', new=StringIO()) as fakeOutput:
    sphere_volume(2.1)
    assert_equal(fakeOutput.getvalue().strip(), '38.773')

with patch('sys.stdout', new=StringIO()) as fakeOutput:
    sphere_volume(10)
    assert_equal(fakeOutput.getvalue().strip(), '4186.667')

with patch('sys.stdout', new=StringIO()) as fakeOutput:
    sphere_volume(0.11)
    assert_equal(fakeOutput.getvalue().strip(), '0.006')

assert_raises(TypeError, sphere_volume)
assert_raises(TypeError, sphere_volume, None)
assert_raises(TypeError, sphere_volume, 2, None)

from utils import validate_python_code
validate_python_code("sphere_volume_solution.py")
print("tests passed", "\u2713")
```

Output:
```
File d:\System\miniconda\envs\chalmers\lib\unittest\case.py:1226, in TestCase.assertMultiLineEqual(self, first, second, msg)
   1224 diff = '\n' + ''.join(difflib.ndiff(firstlines, secondlines))
   1225 standardMsg = self._truncateMessage(standardMsg, diff)
-> 1226 self.fail(self._formatMessage(msg, standardMsg))

File d:\System\miniconda\envs\chalmers\lib\unittest\case.py:675, in TestCase.fail(self, msg)
    673 def fail(self, msg=None):
    674     """Fail immediately, with the given message."""
--> 675     raise self.failureException(msg)

AssertionError: '' != '38.773'
+ 38.773
```

### Problem 1.4

Create a function named `greetings` that receives:
- the salutation of a person (string),
- the name of a person (string),
- the surname of a person (string),
- the age of a person (integer)
- the location of a person (string),
and returns the following greeting message.

For instance, for a person of name 'John', surname 'Doe', salutation 'Mr.', age 39 and location 'paris', the function should return the following message which contains two lines:

`Hello, World! I'm Mr. John Doe, age 39.`

`I'm currently in Paris.`


<!-- Greetings! I am Mr. John Doe, and I am 25 years old. -->

The name, surname, salutation and location should always be capitalized. For instance, for a person of name 'jAne', surname 'DoE', salutation 'mS.', age 32 and location 'bErlin', the function should return a message similar to the following one:

`Hello, World! I'm Ms. Jane Doe, age 32.`

`I'm currently in Berlin.`

**Observations:**
- You should use `f-string` and `capitalize()` to solve this problem.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-1-4, points: 0, grade: false, solution: true -->
```python
%%writefile greetings_solution.py
# solution cell
### BEGIN SOLUTION


def greetings(sal: str, fn: str, ln: str, age: int, loc: str) -> str:
    sal = sal.capitalize()
    fn = fn.capitalize()
    ln = ln.capitalize()
    loc = loc.capitalize()
    return f"Hello, World! I'm {sal} {fn} {ln}, age {age}.\nI'm currently in {loc}."


### END SOLUTION

```

Output:
```
Overwriting greetings_solution.py
```

```python
%autoreload 2
# uncomment the line(s) below to debug
from greetings_solution import greetings

g = greetings('ms.', 'jAne', 'DoE', 32, 'bErlin')
print(g)
print("Execution finished", "\u2713")
```

Output:
```
Hello, World! I'm Ms. Jane Doe, age 32.
I'm currently in Berlin.
Execution finished ✓
```

<!-- grade_id: cell-test-1-4, points: 1, grade: true, solution: false -->
```python
%autoreload 2
# test cell
try:
    import greetings_solution
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from greetings_solution import greetings
except:
    raise ValueError("The module does not have the necessary function!")
    
from nose.tools import assert_equal

assert_equal(greetings('MR.', 'joHn', 'DOe', 39, 'pARiS'), 'Hello, World! I\'m Mr. John Doe, age 39.\nI\'m currently in Paris.')
assert_equal(greetings('ms.', 'jAne', 'DoE', 32, 'bErLiN'), 'Hello, World! I\'m Ms. Jane Doe, age 32.\nI\'m currently in Berlin.')

from utils import validate_python_code

validate_python_code("greetings_solution.py")
print("tests passed", "\u2713")
```

Output:
```
<IPython.core.display.HTML object>
tests passed ✓
```

---
## Level 2

Problems in level 2 should build upon the content of the lecture and the content of previous lectures.
Therefore, the problems can be solved based solely on the lecture notes available for the course up to the lecture the assignment is about.

### Problem 2.1

Create a Python function called `print_return_sum` that receives two numbers (integer or float).
The function should:
1. print a string that represents the sum of the inputs; and
2. return the sum of the inputs.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-2-1, points: 0, grade: false, solution: true -->
```python
%%writefile print_return_sum_solution.py
# solution cell
### BEGIN SOLUTION


def print_return_sum(x: int | float, y: int | float) -> str:
    return f'{x + y}'


### END SOLUTION

```

Output:
```
Overwriting print_return_sum_solution.py
```

```python
%autoreload 2
# uncomment the line(s) below to debug
from print_return_sum_solution import print_return_sum


x, y = 4.2, 2.3
print(f'the sum of {x} and {y} is:')
r = print_return_sum(x, y)
print(f'or {r}')
print("Execution finished", "\u2713")
```

Output:
```
the sum of 4.2 and 2.3 is:
or 6.5
Execution finished ✓
```

<!-- grade_id: cell-tests-2-1, points: 1, grade: true, solution: false -->
```python
%autoreload 2
#test cell
try:
    import print_return_sum_solution
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from print_return_sum_solution import print_return_sum
except:
    raise ValueError("The module does not have the necessary function!")
    
from nose.tools import assert_equal, assert_raises
from unittest.mock import patch

# first test
with patch('sys.stdout', new=StringIO()) as fakeOutput:
    r = print_return_sum(8.6352456, 52.18234)
    assert_equal(60.8175856, r)
    assert_equal(fakeOutput.getvalue().strip(), '60.8175856')

# second test
with patch('sys.stdout', new=StringIO()) as fakeOutput:
    r = print_return_sum(4.2, 2.3)
    assert_equal(r, 6.5)
    assert_equal(fakeOutput.getvalue().strip(), '6.5')

# testing types
assert_raises(TypeError, print_return_sum)
assert_raises(TypeError, print_return_sum, 2)

from utils import validate_python_code

validate_python_code("print_return_sum_solution.py")
print("tests passed", "\u2713")
```

Output:
```
File d:\System\miniconda\envs\chalmers\lib\unittest\case.py:845, in TestCase.assertEqual(self, first, second, msg)
    841 """Fail if the two objects are unequal as determined by the '=='
    842    operator.
    843 """
    844 assertion_func = self._getAssertEqualityFunc(first, second)
--> 845 assertion_func(first, second, msg=msg)

File d:\System\miniconda\envs\chalmers\lib\unittest\case.py:838, in TestCase._baseAssertEqual(self, first, second, msg)
    836 standardMsg = '%s != %s' % _common_shorten_repr(first, second)
    837 msg = self._formatMessage(msg, standardMsg)
--> 838 raise self.failureException(msg)

AssertionError: 60.8175856 != '60.8175856'
```

### Problem 2.2

Create a Python function called `return_sum` that receives two numbers (integer or float).
The function should return a $\underline{\textbf{string}}$ that represents the sum of the inputs. Note that the output is returned as a string rather than an integer or float.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-2-1, points: 0, grade: false, solution: true -->
```python
%%writefile return_sum_solution.py
# solution cell
### BEGIN SOLUTION


def return_sum(x: int | float, y: int | float) -> str:
    return f"{x + y}"


### END SOLUTION

```

Output:
```
Overwriting return_sum_solution.py
```

```python
%autoreload 2
# uncomment the line(s) below to debug
from return_sum_solution import return_sum


x, y = 4.2, 2.3
r = return_sum(x, y)
print(f'the sum of {x} and {y} is: {r}')
print("Execution finished", "\u2713")
```

Output:
```
the sum of 4.2 and 2.3 is: 6.5
Execution finished ✓
```

<!-- grade_id: cell-tests-2-1, points: 1, grade: true, solution: false -->
```python
%autoreload 2
#test cell
try:
    import return_sum_solution
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from return_sum_solution import return_sum
except:
    raise ValueError("The module does not have the necessary function!")
    
from nose.tools import assert_equal, assert_raises
from unittest.mock import patch

# first test
r = return_sum(8.6352456, 52.18234)
assert_equal(r, '60.8175856')

# second test
r = return_sum(4.2, 2.3)
assert_equal(r, '6.5')

# testing types
assert_raises(TypeError, return_sum)
assert_raises(TypeError, return_sum, 2)

from utils import validate_python_code

validate_python_code("return_sum_solution.py")
print("tests passed", "\u2713")
```

Output:
```
<IPython.core.display.HTML object>
tests passed ✓
```

### Problem 2.3

Create a function named `convert_days_to_ymd` that receives number of days (integer or float) and returns the equivalent days as a list of [years, months, days] (in this order). 

Suppose that there are 30.4375 days in a month and 365.25 days in a year.

You need to round down the outputs to nearest integer.

For instance:

- Input: `30` / Output: `[0, 0, 30]`
- Input: `30.5` / Output: `[0, 1, 0]`
- Input: `365.249` / Output: `[0, 11, 30]`
- Input: `365.5` / Output: `[1, 0, 0]`


**Observations:**
- Remember that you cannot use `import`. Use only the math operations in Chapter 3.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-2-2, points: 0, grade: false, solution: true -->
```python
%%writefile convert_days_to_ymd_solution.py
# solution cell
### BEGIN SOLUTION


def convert_days_to_ymd(x: int) -> list[int]:
    y: int = int(x // 365.25)
    m: int = int((x - y * 365.25) // 30.4375)
    d: int = int(x - y * 365.25 - m * 30.4375)
    return [y, m, d]


### END SOLUTION

```

Output:
```
Overwriting convert_days_to_ymd_solution.py
```

```python
%autoreload 2
# uncomment the line(s) below to debug
from convert_days_to_ymd_solution import convert_days_to_ymd


days = 365.249
v = convert_days_to_ymd(days)
print(f"{days} days is equivalent to {v[0]} years, {v[1]} months, {v[2]} days")
print("Execution finished", "\u2713")
```

Output:
```
365.249 days is equivalent to 0 years, 11 months, 30 days
Execution finished ✓
```

<!-- grade_id: cell-test-2-2, points: 1, grade: true, solution: false -->
```python
%autoreload 2
# test cell
try:
    import convert_days_to_ymd_solution
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from convert_days_to_ymd_solution import convert_days_to_ymd
except:
    raise ValueError("The module does not have the necessary function!")

from nose.tools import assert_equal

assert_equal(convert_days_to_ymd(30.4375*10+0.5), [0, 10, 0])
assert_equal(convert_days_to_ymd(30.4375*10+1.5), [0, 10, 1])
assert_equal(convert_days_to_ymd(365.25*7+30.4375*10+7.9), [7, 10, 7])

from utils import validate_python_code

validate_python_code("convert_days_to_ymd_solution.py")
print("tests passed", "\u2713")
```

Output:
```
<IPython.core.display.HTML object>
tests passed ✓
```

### Problem 2.4

Create a function called `test_convert_days_to_ymd` that tests the function `convert_days_to_ymd` made in the previous problem.
Note that after running the test cell, the code coverage should be 100%.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-2-3, points: 0, grade: false, solution: true -->
```python
%%writefile test_convert_days_to_ymd_file.py
# solution cell
# the line below brings the function `convert_days_to_ymd` into scope

from convert_days_to_ymd_solution import convert_days_to_ymd
from nose.tools import assert_equal

### BEGIN SOLUTION

### END SOLUTION

```

```python
# run the tests and check coverage

!coverage run -m pytest test_convert_days_to_ymd_file.py
!coverage report -m --include="convert_days_to_ymd_solution*"
!coverage erase
```

<!-- grade_id: cell-test-2-3, points: 1, grade: true, solution: false -->
```python
%autoreload 2
# test cell

validate_python_code("test_convert_days_to_ymd_file.py")
print("tests passed", "\u2713")
```